Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `???`, `YOUR CODE HERE`, "???", "YOUR ANSWER HERE", as well as your name and collaborators below:

```scala
val NAME = ""
val COLLABORATORS = ""
```

---

# Exercise: Higher-Order Functions

<!-- 3 Expressions -->

<!-- 4 Binding and Scope -->

<!-- 8 Recursion -->

<!-- 9 Inductive Data Types -->

<!-- 11 Concrete Syntax -->

<!-- 12 Abstract Syntax and Parsing -->

<!-- 13 Exercise: Syntax -->

<!-- 14 Static Scoping -->

<!-- 15 Judgments -->

<!-- 16 Variables, Basic Values, and Judgments Lab -->

<!-- 18 Operational Semantics -->

<!-- 19 Functions and Dynamic Scoping -->

<!-- 20 Big-Step Exercise -->

<!-- 21 Evaluation Order  -->

### Learning Goals

The primary learning goal of this exercise is to get experience
programming with higher-order functions.

### Instructions

This assignment asks you to write Scala code. There are restrictions
associated with how you can solve these problems. Please pay careful
heed to those. If you are unsure, ask the course staff.

Note that `???` indicates that there is a missing function or code
fragment that needs to be filled in. Make sure that you remove the `???`
and replace it with the answer.

Use the test cases provided to test your implementations. You are also
encouraged to write your own test cases to help debug your work.
However, please delete any extra cells you may have created lest they
break an autograder.

### Imports

```scala
// Run this cell FIRST before testing.
import $ivy.`org.scalatest::scalatest:3.2.19`, org.scalatest._, events._, flatspec._
def report(suite: Suite): Unit = suite.execute(stats = true)
def assertPassed(suite: Suite): Unit =
  suite.run(None, Args(new Reporter {
    def apply(e: Event) = e match {
      case e @ (_: TestFailed) => assert(false, s"${e.message} (${e.testName})")
      case _ => ()
    }
  }))
def passed(points: Int): Unit = {
  require(points >=0)
  if (points == 1) println("*** 🎉 Tests Passed (1 point) ***")
  else println(s"*** 🎉 Tests Passed ($points points) ***")
}
def test(suite: Suite, points: Int): Unit = {
  report(suite)
  assertPassed(suite)
  passed(points)
}
```

## Collections

When working with and organizing data, we primarily use collections from
Scala’s standard library. One of the most fundamental operations that
one needs to perform with a collection is to iterate over the elements.
Like many other languages with first-class functions (e.g., Python, ML),
the Scala library provides various iteration operations via
*higher-order functions*. Higher-order functions are functions that take
functions as parameters. The function parameters are often called
*callbacks*, and for collections, they typically specify what the
library client wants to do for each element. We have seen examples of
these functions in class. In the following questions, we practice both
writing such higher-order functions in a library and using them as a
client.

### Lists

First, we will implement functions that eliminate consecutive duplicates
of list elements. If a list contains repeated elements they should be
replaced with a single copy of the element. The order of the elements
should not be changed. For example, the following
`List(1, 2, 2, 3, 3, 3)` should be converted to `List(1,2,3)`.

<span class="theorem-title">**Exercise 1 (5 points)**</span> Write a
function `compressRec` that implements this behavior. Implement the
function by direct recursion (e.g., pattern match on `l` and call
`compressRec` recursively). Do not call any `List` library methods.

**Edit this cell:**

<!-- grade_id: compress-rec-answer, points: 0, grade: false, solution: true -->
```scala
def compressRec[A](l: List[A]): List[A] = l match {
  case Nil | _ :: Nil => 
    ???
  case h1 :: (t1 @ (h2 :: _)) => 
    ???
}
```

#### Notes

-   This exercise is from [Ninety-Nine Scala
    Problems](http://aperiodic.net/phil/scala/s-99/). Some sample
    solutions are given there, which you are welcome to view. However,
    we *strongly* recommend you attempt this exercise before looking
    there. The purpose of the exercise is to get some practice for the
    later part of this homework. Note that the solutions there do not
    satisfy the requirements here, as they use library functions. If at
    some point you feel like you need more practice with collections,
    this site is a good resource.

#### Tests

<!-- grade_id: eval-compress-rec-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  val l1 = List(1, 2, 2, 3, 3, 3)
  val r1 = List(1, 2, 3)

  "compressRec" should "Test Case 1" in { assertResult( r1 ) { compressRec(l1) } }
}, 5)
```

<span class="theorem-title">**Exercise 2 (5 points)**</span> Write a
different function `compressFold` that re-implements the behavior of
`compressRec` using the `foldRight` method from the `List` library. The
call to `foldRight` has been provided for you. Do not call
`compressFold` recursively or any other `List` library methods.

**Edit this cell:**

<!-- grade_id: compress-fold-answer, points: 0, grade: false, solution: true -->
```scala
def compressFold[A](l: List[A]): List[A] = l.foldRight(Nil: List[A]){ (h, acc) => 
  ???
}
```

#### Tests

<!-- grade_id: eval-compress-fold-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  val l1 = List(1, 2, 2, 3, 3, 3)
  val r1 = List(1, 2, 3)

  "compressFold" should "Test Case 1" in { assertResult( r1 ) { compressFold(l1) } }
}, 5)
```

<span class="theorem-title">**Exercise 3 (5 points)**</span> Explain in
1–2 sentences the similarities and differences between your two
implementations: `compressRec` and `compressFold`.

**Edit this cell:**

<!-- grade_id: explain-compress, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

<span class="theorem-title">**Exercise 4 (5 points)**</span> Implement a
higher-order recursive function `mapFirst` that finds the first element
of `l: List[A]` where `f: A => Option[A]` applied to it returns
`Some(a)` for some value `a`. The function should replace that element
with `a` and leave `l` the same everywhere else. For example,

``` scala
mapFirst(List(1,2,-3,4,-5)) { i => if (i < 0) Some(-i) else None } 
```

should result in `List[Int] = List(1, 2, 3, 4, -5)`.

**Edit this cell:**

<!-- grade_id: map-first-answer, points: 0, grade: false, solution: true -->
```scala
def mapFirst[A](l: List[A])(f: A => Option[A]): List[A] = l match {
  case Nil => 
    ???
  case h :: t => 
    ???
}
```

#### Tests

<!-- grade_id: eval-map-first-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  val l1 = List(1, 2, -3, 4, -5)
  val r1 = List(1, 2, 3, 4, -5)

  "mapFirst" should "Test Case 1" in { 
    assertResult(r1) {
      mapFirst(l1) { (i: Int) => if (i < 0) Some(-i) else None }
    }
  }
}, 5)
```

<span class="theorem-title">**Exercise 5 (5 points)**</span> Write a
function `composeMap` that sequentially applies a list of functions of
type `A => A` to all the elements of a `List[A]`. For example, if we
have a list of functions `List(f3, f2, f1)`, and a list `List(a, b)`, we
want to output `List(f3(f2(f1(a))), f3(f2(f1(b))))`.

**Edit this cell:**

<!-- grade_id: compose-map-answer, points: 0, grade: false, solution: true -->
```scala
def composeMap[A](functions: List[A => A])(l: List[A]): List[A] =
  ???
```

#### Tests

<!-- grade_id: eval-compose-map-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  def plusOne(x: Int): Int = x + 1
  def square(x: Int): Int = x * x
  def timesThree(x: Int): Int = x * 3

  val l1 = List(1, 2, 3)
  val r1 = List(10, 37, 82)
  val fs: List[Int => Int] = List(plusOne, square, timesThree)

  "composeMap" should "Test Case 1" in { assertResult(r1) { composeMap(fs)(l1) } }
}, 5)
```

### Maps

Recall the `Map[A, B]` data structure from class. Also, recall the
higher-order function `map`. To avoid confusion we will use the upper
case `Map` to refer to the type, and the lowercase `map` to refer to the
higher-order function.

<span class="theorem-title">**Exercise 6 (5 points)**</span> Implement a
function `mapValues` that takes a `Map[A,B]` and a callback function
`f: B => C`, that applies `f` to all the values in the `Map`. Your
function should use the Scala collections `map` method. Do not use the
standard library method `mapValues` on `Map` (however, note that the
behavior of your function should be exactly the same as the `mapValues`
standard library method).

**Edit this cell:**

<!-- grade_id: update-map-answer, points: 0, grade: false, solution: true -->
```scala
def mapValues[A, B, C](m: Map[A, B])(f: B => C): Map[A, C] = 
  ???
```

#### Tests

<!-- grade_id: eval-update-map-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  val m1 = Map("a" -> 1, "b" -> 2, "c" -> 3)
  val r1 = Map("a" -> 1, "b" -> 4, "c" -> 9)

  "mapValues" should "Test Case 1" in { 
    assertResult(r1) {
      mapValues(m1) { (i: Int) => i * i }
    }
  }
}, 5)
```

<span class="theorem-title">**Exercise 7 (5 points)**</span> As we
mentioned above, we just reimplemented `mapValues` using the `map`
method from the standard library. Earlier, we implemented higher-order
functions on `List`s recursively. Could we have implemented `mapValues`
on `Map`s recursively? If yes, give an example implementation. If no,
explain why we cannot, and what makes `Map` different from `List` in
this case.

**Edit this cell:**

<!-- grade_id: explain-map, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

### Trees

Recall the binary tree data type:

```scala
sealed trait Tree
case object Empty extends Tree
case class Node(l: Tree, d: Int, r: Tree) extends Tree
```

<span class="theorem-title">**Exercise 8 (10 points)**</span> Implement
a higher-order function `foldLeft` that performs an in-order traversal
of the input `t: Tree`, calling the callback `f` starting with `z` to
accumulate a result. For example, suppose the in-order traversal of the
input tree `t` yields the sequence of data values: $d_1,d_2,\dots,d_n$.
Then, `foldLeft(t)(z)(f)` yields
$f(f(\dots(f(f(z,d_1),d_2))\dots d_{n-1}),d_n)$

**Edit this cell:**

<!-- grade_id: tree-foldleft-answer, points: 0, grade: false, solution: true -->
```scala
def foldLeft[A](t: Tree)(z: A)(f: (A, Int) => A): A = {
  def loop(acc: A, t: Tree): A = t match {
    case Empty => 
      ???
    case Node(l, d, r) => 
      ???
  }
  ???
}
```

We have provided a test client `sum` that computes the sum of all of the
data values in the tree using your `foldLeft` method, along with some
helper functions to build trees more easily. Feel free to use them to
write more test cases for your code.

```scala
// An example use of foldLeft
def sum(t: Tree): Int = foldLeft(t)(0){ (acc, d) => acc + d }

// In-order insertion into a binary search tree
def insert(t: Tree, n: Int): Tree = t match {
  case Empty => Node(Empty, n, Empty)
  case Node(l, d, r) =>
    if (n < d) Node(insert(l, n), d, r) else Node(l, d, insert(r, n))
}

// Create a tree from a list. An example use of the List.foldLeft method.
def treeFromList(l: List[Int]): Tree =
  l.foldLeft(Empty: Tree){ (acc, i) => insert(acc, i) }
```

#### Tests

<!-- grade_id: eval-fold-left-tester, points: 10, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  val t1 = treeFromList(List(1, 2, 3))
  "foldLeft" should "Test Case 1" in { assertResult(6) { sum(t1) } }

  val l2 = List(3, 4, 8, 2, 1, 7, 10)
  val sortl2 = l2.sorted
  val t2 = treeFromList(l2)

  def rev(acc: List[Int], x: Int): List[Int] = x :: acc
  def prod(acc: Int, x: Int): Int = x * acc

  it should "Test Case 2" in {
    assertResult( sortl2.foldLeft(Nil: List[Int])(rev) ) {
      foldLeft(t2)(Nil: List[Int])(rev)
    }
  }

  it should "Test Case 3" in {
    assertResult( l2.foldLeft(1)(prod) ) {
      foldLeft(t2)(1)(prod)
    }
  }

}, 10)
```

<span class="theorem-title">**Exercise 9 (10 points)**</span> Using your
`foldLeft` function, write a client function `strictlyOrdered` that
checks if the data values of an in-order traversal of `t` are in
strictly ascending order. For example, suppose the in-order traversal of
the input tree `t` yields the sequence of data values:
$d_1,d_2,\dots,d_n$, the `strictlyOrdered` should return `true` iff
$d_1 < d_2 < \dots < d_n$.

**Edit this cell:**

<!-- grade_id: tree-strictly-ordered-answer, points: 0, grade: false, solution: true -->
```scala
def strictlyOrdered(t: Tree): Boolean = {
  val (b, _) = foldLeft(t)((true, None: Option[Int])) {
    ???
  }
  b
}
```

#### Tests

<!-- grade_id: eval-strictly-ordered-tester, points: 10, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  "strictlyOrdered" should "Test Case 1" in { assert(!strictlyOrdered(treeFromList(List(1,1,2)))) }
  it should "Test Case 2" in { assert(strictlyOrdered(treeFromList(List(1,2)))) }
  it should "Test Case 3" in { assert(!strictlyOrdered(treeFromList(List(-1,0,1,1,2)))) }
  it should "Test Case 4" in { assert(!strictlyOrdered(treeFromList(List(2,1,2,3,4)))) }
  it should "Test Case 5" in { assert(strictlyOrdered(treeFromList(List(1,2,3,4,5)))) }
}, 10)
```

Now, we will write a higher-order function that uses our recent
higher-order functions as a client.

<span class="theorem-title">**Exercise 10 (5 points)**</span> Implement
a function `foldLeftTrees` that take as input a `lt: List[Tree]`, and
applies a callback `f` to all of their nodes in-order (of both the
`List` and the nested `Tree`s) starting with initial value `z` to
accumulate a result.

For example, calling `foldLeftTrees` on

``` scala
val lt = 
  List(
    Node(Node(Empty, 1, Empty), 2, Node(Empty, 3, Empty)),
    Node(Node(Empty, 4, Empty), 5, Node(Empty, 6, Empty))
  )
```

<pre><span class="ansi-cyan-fg">lt</span>: <span class="ansi-green-fg">List</span>[<span class="ansi-green-fg">Node</span>] = <span class="ansi-yellow-fg">List</span>(
  <span class="ansi-yellow-fg">Node</span>(
    l = <span class="ansi-yellow-fg">Node</span>(l = Empty, d = <span class="ansi-green-fg">1</span>, r = Empty),
    d = <span class="ansi-green-fg">2</span>,
    r = <span class="ansi-yellow-fg">Node</span>(l = Empty, d = <span class="ansi-green-fg">3</span>, r = Empty)
  ),
  <span class="ansi-yellow-fg">Node</span>(
    l = <span class="ansi-yellow-fg">Node</span>(l = Empty, d = <span class="ansi-green-fg">4</span>, r = Empty),
    d = <span class="ansi-green-fg">5</span>,
    r = <span class="ansi-yellow-fg">Node</span>(l = Empty, d = <span class="ansi-green-fg">6</span>, r = Empty)
  )
)</pre>

and a function that sums integers should return `21`.

Use only `foldLeft` on `Tree` that you have defined above and `foldLeft`
on `List[Tree]` provided in the Scala standard library.

**Edit this cell:**

<!-- grade_id: reduce-trees-answer, points: 0, grade: false, solution: true -->
```scala
def foldLeftTrees[B](lt: List[Tree])(z: B)(f: (B, Int) => B): B =
  ???
```

#### Tests

<!-- grade_id: eval-reduce-trees-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  def sum(trees: List[Tree]): Int = foldLeftTrees(trees)(0){ (acc, d) => acc + d }

  val l1 = List(1, 2, 3)
  val l2 = List(4, 5, 6)
  val t1 = treeFromList(l1)
  val t2 = treeFromList(l2)
  val lt1 = List(t1, t2)

  val l1l2 = l1 ::: l2

  "foldLeftTrees" should "Test Case 1" in { assertResult(21){ sum(lt1) } }

  it should "reverse and flatten" in {
    assertResult( l1l2.reverse ) { foldLeftTrees(lt1)(Nil: List[Int]){ (acc, i) => i :: acc } }
  }
}, 5)
```

## flatMap

<span class="theorem-title">**Exercise 11 (5 points)**</span> Recall the
`flatMap` function from class on `List`. Write a new function
`flatMapNoRec` that implements the same behavior without direct
recursion; instead, use `foldRight` and `:::`.

**Edit this cell:**

<!-- grade_id: flat-map-no-rec-answer, points: 0, grade: false, solution: true -->
```scala
def flatMapNoRec[A, B](l: List[A])(f: A => List[B]): List[B] = 
  ???
```

#### Tests

<!-- grade_id: eval-flat-map-no-rec-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  def duplicate[A](l: List[A]) = flatMapNoRec(l) { a => List(a, a) }

  val l1 = List(1, 2, 3)
  val r1 = List(1, 1, 2, 2, 3, 3)

  "flatMapNoRec" should "Test Case 1" in { assertResult(r1) { duplicate(l1) } }
}, 5)
```

Now, we will try to use `flatMap` to do something useful.

<span class="theorem-title">**Exercise 12 (5 points)**</span> Write a
function `getAllWords` that takes in a `List[String]` of sentences and
returns a `List[String]` of all the words in the sentences. For
simplicity, our sentences will have no punctuation, and all words will
be separated by a single space. For example, given the input
`List("I love 3155", "Anish is the best TA")`, the `getAllWords`
function should return
`List("I", "love", "3155", "Anish", "is", "the", "best", "TA")`.

Use the String method `split` to separate a sentence into its component
words:

``` scala
"Functions are values".split(" ").toList
```

<pre><span class="ansi-cyan-fg">res22</span>: <span class="ansi-green-fg">List</span>[<span class="ansi-green-fg">String</span>] = <span class="ansi-yellow-fg">List</span>(<span class="ansi-green-fg">"Functions"</span>, <span class="ansi-green-fg">"are"</span>, <span class="ansi-green-fg">"values"</span>)</pre>

**Edit this cell:**

<!-- grade_id: get-all-words-answer, points: 0, grade: false, solution: true -->
```scala
def getAllWords(sentences: List[String]): List[String] = 
  ???
```

#### Tests

<!-- grade_id: eval-get-all-words-tester, points: 5, grade: true, solution: false -->
```scala
test(new AnyFlatSpec {
  val l1 = List("I love 3155", "Anish is the best TA")
  val r1 = List("I", "love", "3155", "Anish", "is", "the", "best", "TA")

  "getAllWords" should "Test Case 1" in { assertResult(r1) { getAllWords(l1) } }
}, 5)
```